rule_files:
  - capi.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machine_status_phase{cluster="galaxy", name="galaxy-72jq5", namespace="giantswarm", phase="Running"}'
        values: "1+0x10 0+0x20" #machine was in Running for 10 minutes and than in
      - series: 'capi_machine_status_phase{cluster="galaxy", name="galaxy-72jq5", namespace="giantswarm", phase="Failed"}'
        values: "0+0x10 1+0x20"
    alert_rule_test:
      - alertname: MachineUnhealthyPhase
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: rocket
              topic: managementcluster
              cluster: galaxy
              name: galaxy-72jq5
              namespace: giantswarm
              phase: Failed
            exp_annotations:
              description: "Machine giantswarm/galaxy-72jq5 stuck in phase Failed for more than 15 minutes."
