rule_files:
  - capi-machinepool.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machinepool_status_phase{phase="Failed", cluster_name="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinepool_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-72r5c", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: MachinePoolIsNotHealthy
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: page
              phase: Failed
              team: phoenix
              topic: managementcluster
              cluster_name: clippaxy
              name: clippaxy-def00
              exported_namespace: giantswarm
            exp_annotations:
              description: "The clusters clippaxy machinepool giantswarm/clippaxy-def00 is not healthy."
              opsrecipe: capi-machinepool/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
      - alertname: MachinePoolPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-72r5c
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "The clusters grumpy machinepool giantswarm/grumpy-72r5c is paused."
              opsrecipe: capi-machinepool/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
