rule_files:
  - capi-machineset.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machineset_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-def99", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: MachineSetPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-def99
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "Machineset giantswarm/grumpy-def99 is paused."
              opsrecipe: capi-machineset/
              dashboard: bdi7iswg81czkcasd/capi-agregated-error-logs-for-capi-controllers
