rule_files:
  - capi-machinedeployment.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machinedeployment_status_phase{phase="Failed", cluster_name="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinedeployment_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-def99", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: MachineDeploymentIsNotHealthy
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              phase: Failed
              team: phoenix
              topic: managementcluster
              cluster_name: clippaxy
              name: clippaxy-def00
              exported_namespace: giantswarm
            exp_annotations:
              description: "The clusters clippaxy machinedeployment giantswarm/clippaxy-def00 is not healthy."
              runbook_url: capi-machinedeployment/
              dashboardUid: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
      - alertname: MachineDeploymentPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-def99
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "The clusters grumpy machinedeployment giantswarm/grumpy-def99 is paused."
              runbook_url: capi-machinedeployment/
              dashboardUid: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
