rule_files:
  - capi-machine.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machine_status_phase{cluster_name="clippaxy", name="clippaxy-72jq5", exported_namespace="giantswarm", phase="Running"}'
        values: "1+0x10 0+0x35"
      - series: 'capi_machine_status_phase{cluster_name="clippaxy", name="clippaxy-72jq5", exported_namespace="giantswarm", phase="Failed"}'
        values: "0+0x10 1+0x35"
      - series: 'capi_machine_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-72r5c", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: MachineUnhealthyPhase
        eval_time: 45m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: page
              team: phoenix
              topic: managementcluster
              cluster_name: clippaxy
              name: clippaxy-72jq5
              exported_namespace: giantswarm
              phase: Failed
            exp_annotations:
              description: "Machine giantswarm/clippaxy-72jq5 stuck in phase Failed for more than 30 minutes."
              opsrecipe: capi-machine/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
      - alertname: MachinePaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-72r5c
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "Machine giantswarm/grumpy-72r5c is paused."
              opsrecipe: capi-machine/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
