rule_files:
  - capi-cluster.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_cluster_status_phase{name="clippaxy", exported_namespace="giantswarm", phase="Provisioned"}'
        values: "1+0x75"
      - series: 'capi_cluster_status_phase{name="clippaxy", exported_namespace="giantswarm", phase="Pending"}'
        values: "1+0x75"
      - series: 'capi_cluster_status_condition{name="grumpy", exported_namespace="giantswarm", status="False", type="Ready"}'
        values: "0+0x10 0+1x65"
      - series: 'capi_cluster_status_condition{name="grumpy", exported_namespace="giantswarm", status="True", type="Ready"}'
        values: "0+1x10 0+0x65"
      - series: 'capi_cluster_annotation_paused{name="grumpy", exported_namespace="giantswarm", paused_value="true"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: ClusterUnhealthyPhase
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: page
              team: phoenix
              topic: managementcluster
              name: clippaxy
              exported_namespace: giantswarm
              phase: Pending
            exp_annotations:
              description: "Cluster giantswarm/clippaxy stuck in Pending phase."
              opsrecipe: capi-cluster/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
      - alertname: ClusterStatusNotReady
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              name: grumpy
              exported_namespace: giantswarm
              status: "False"
              type: Ready
            exp_annotations:
              description: "Cluster giantswarm/grumpy is not ready."
              opsrecipe: capi-cluster/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
      - alertname: ClusterPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              name: grumpy
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "The cluster giantswarm/grumpy is paused."
              opsrecipe: capi-cluster/
              dashboard: bdi7iswg81czkcasd/capi-aggregated-error-logs-for-capi-controllers?orgId=2
