rule_files:
  - capi-machinepool.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machinepool_status_phase{phase="Failed", cluster_id="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_cluster_info{cluster_id="clippaxy", provider="capa"}'
        values: "1+0x75"
      - series: 'capi_machinepool_annotation_paused{paused_value="true",cluster_id="grumpy", name="grumpy-72r5c", exported_namespace="giantswarm"}'
        values: "0+1x75"
      - series: 'capi_cluster_info{cluster_id="grumpy", provider="capa"}'
        values: "1+0x75"
    alert_rule_test:
      - alertname: MachinePoolIsNotHealthy
        eval_time: 25m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              provider: capa
              severity: page
              phase: Failed
              team: phoenix
              topic: managementcluster
              cluster_id: clippaxy
              name: clippaxy-def00
              exported_namespace: giantswarm
            exp_annotations:
              description: "The clusters clippaxy machinepool giantswarm/clippaxy-def00 is not healthy."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/capi-machinepool/
              __dashboardUid__: bdi7iswg81czkcasd
              dashboardQueryParams: "orgId=2"
      - alertname: MachinePoolPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              provider: capa
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_id: grumpy
              name: grumpy-72r5c
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "The clusters grumpy machinepool giantswarm/grumpy-72r5c is paused."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/capi-machinepool/
              __dashboardUid__: bdi7iswg81czkcasd
              dashboardQueryParams: "orgId=2"
