rule_files:
  - capi-cluster.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_cluster_status_phase{cluster_id="clippaxy", exported_namespace="giantswarm", phase="Pending"}'
        values: "1x75"
      - series: 'capi_cluster_info{cluster_id="clippaxy", provider="capa"}'
        values: "1x75"
      - series: 'capi_cluster_status_condition{cluster_id="grumpy", exported_namespace="giantswarm", status="False", type="Ready"}'
        values: "0+0x10 0+1x65"
      - series: 'capi_cluster_status_condition{cluster_id="grumpy", exported_namespace="giantswarm", status="True", type="Ready"}'
        values: "0+1x10 0+0x65"
      - series: 'capi_cluster_annotation_paused{cluster_id="grumpy", exported_namespace="giantswarm", paused_value="true"}'
        values: "0+1x75"
      - series: 'capi_cluster_info{cluster_id="grumpy", provider="capa"}'
        values: "1x75"
    alert_rule_test:
      - alertname: ClusterUnhealthyPhase
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: clippaxy
              severity: page
              team: phoenix
              topic: managementcluster
              exported_namespace: giantswarm
              phase: Pending
              provider: capa
            exp_annotations:
              description: "Cluster giantswarm/clippaxy stuck in Pending phase."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/capi-cluster/
              __dashboardUid__: bdi7iswg81czkcasd
              dashboardQueryParams: "orgId=2"
      - alertname: ClusterStatusNotReady
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: grumpy
              provider: capa
              severity: notify
              team: phoenix
              topic: managementcluster
              exported_namespace: giantswarm
              status: "False"
              type: Ready
            exp_annotations:
              description: "Cluster giantswarm/grumpy is not ready."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/capi-cluster/
              __dashboardUid__: bdi7iswg81czkcasd
              dashboardQueryParams: "orgId=2"
      - alertname: ClusterPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: grumpy
              provider: capa
              severity: notify
              team: phoenix
              topic: managementcluster
              exported_namespace: giantswarm
              paused_value: "true"
            exp_annotations:
              description: "The cluster giantswarm/grumpy is paused."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/capi-cluster/
              __dashboardUid__: bdi7iswg81czkcasd
              dashboardQueryParams: "orgId=2"

  - interval: 1m
    input_series:
      # Taken from a real upgrade failure, unnecessary labels removed
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-6d9q7", namespace="giantswarm", provider_id="aws:///eu-central-1c/i-0643d384aa5d42d57"}
        values: "_x26 1x10 _x60"
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-7npqj", namespace="giantswarm", provider_id="aws:///eu-central-1b/i-0325c255e1a65cd07"}
        values: "_x47 1x10 _x60"
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-9rhg7", namespace="giantswarm", provider_id="aws:///eu-central-1b/i-02636d5d4c6eca6e4"}
        values: "_x16 1x8 _x60"
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-hw757", namespace="giantswarm", provider_id="aws:///eu-central-1a/i-081d154aba0f21660"}
        values: "1x60"

      # This `name` appeared twice with different labels (e.g. with/without `provider_id`).
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-dhd55", namespace="giantswarm", provider_id="aws:///eu-central-1a/i-028756797a3e68e09"}
        values: "_x36 1x7 _x60"
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-dhd55", namespace="giantswarm"}
        values: "_x34 1x1 _x60"

      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-q49zf", namespace="giantswarm", provider_id="aws:///eu-central-1b/i-0a026aff280d4cb18"}
        values: "1x60"
      - series: capi_machine_info{cluster_id="oopsie", control_plane_name="oopsie", name="oopsie-sf9zr", namespace="giantswarm", provider_id="aws:///eu-central-1c/i-025ad92e96fa3a2c9"}
        values: "1x60"


      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-6d9q7", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "_x33 1x3 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-6d9q7", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "_x33 0x3 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-7npqj", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "_x55 1x3 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-7npqj", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "_x55 1x3 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-9rhg7", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "_x24 1x1 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-9rhg7", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "_x24 0x1 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-dhd55", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "_x44 1x3 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-dhd55", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "_x44 0x3 _x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-hw757", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "0x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-hw757", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "0x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-q49zf", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "0x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-q49zf", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "0x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-sf9zr", namespace="giantswarm", status="False", type="HealthCheckSucceeded"}
        values: "0x60"
      - series: capi_machine_status_condition{cluster_id="oopsie", name="oopsie-sf9zr", namespace="giantswarm", status="Unknown", type="HealthCheckSucceeded"}
        values: "0x60"
    alert_rule_test:
      - alertname: ClusterControlPlaneMachineStatusNotHealthy
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_monitoring_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: oopsie
              namespace: giantswarm
              severity: notify
              team: phoenix
              topic: managementcluster
            exp_annotations:
              description: "Cluster 'giantswarm/oopsie' has one or more non-ready control plane nodes. Failing cluster upgrade/node-roll?"
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/capi-cluster/
