---
rule_files:
  - certificate.management-cluster.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'cert_exporter_secret_not_after{cluster_id="gauss", cluster_type="management_cluster", secretkey="tls.crt", certificatename="aws-admission-controller-certificates", exported_namespace="giantswarm", provider="capa"}'
        values: "1x20 1x20 0+0x20"
      - series: 'cert_exporter_certificate_cr_not_after{cluster_id="gauss", cluster_type="management_cluster", name="aws-admission-controller-certificates", exported_namespace="giantswarm", provider="capa"}'
        values: "1x20 _x20 0+0x20"
    alert_rule_test:
      - alertname: ManagementClusterCertificateIsMissing
        eval_time: 15m
      - alertname: ManagementClusterCertificateIsMissing
        eval_time: 35m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "true"
              certificatename: aws-admission-controller-certificates
              cluster_id: gauss
              exported_namespace: giantswarm
              provider: capa
              severity: page
              team: phoenix
              topic: security
            exp_annotations:
              description: 'Cannot renew Certificate for Secret giantswarm/aws-admission-controller-certificates on gauss because it is missing.'
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/managed-app-cert-manager/missing-certificate-for-secret/
      - alertname: ManagementClusterCertificateIsMissing
        eval_time: 55m
