---
rule_files:
- mimir-to-grafana-cloud-exporter.rules.yml

tests:
  # Tests for `MimirToGrafanaCloudExporterDown` alert
  - interval: 1m
    input_series:
      - series: 'up{job="mimir/mimir-to-grafana-cloud", cluster_id="golem", installation="golem", namespace="mimir", customer="giantswarm", pipeline="testing", provider="capa", region="eu-west-2"}'
        values: "_x60 1+0x60 0+0x60 1+0x60"
    alert_rule_test:
      - alertname: MimirToGrafanaCloudExporterDown
        eval_time: 70m
      - alertname: MimirToGrafanaCloudExporterDown
        eval_time: 160m
        exp_alerts:
          - exp_labels:
              area: platform
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "golem"
              customer: "giantswarm"
              installation: "golem"
              namespace: "mimir"
              job: "mimir/mimir-to-grafana-cloud"
              pipeline: "testing"
              provider: "capa"
              region: "eu-west-2"
            exp_annotations:
              dashboard: "iWowmlSmk/prometheus?var-cluster=mimir-to-grafana-cloud"
              description: "Prometheus Mimir to Grafana-Cloud is down."
              opsrecipe: "mimir-grafana-cloud-exporter-failing/"
      - alertname: MimirToGrafanaCloudExporterDown
        eval_time: 200m
  # Tests for `MimirToGrafanaCloudExporterFailures` alert
  - interval: 1m
    input_series:
      # remote read is working for 2 hours and then fails for 1 hour
      - series: 'prometheus_remote_storage_read_queries_total{job="mimir/mimir-to-grafana-cloud", cluster_id="golem", customer="giantswarm", installation="golem", namespace="mimir", pipeline="testing", provider="capa", region="eu-west-2"}'
        values: "_x60 0+10x60 0+0x60 0+10x180"
      # remote write has no failure for 4 hours and then fails for 2 hours
      - series: 'prometheus_remote_storage_samples_failed_total{job="mimir/mimir-to-grafana-cloud", cluster_id="golem", customer="giantswarm", installation="golem", namespace="mimir", pipeline="testing", provider="capa", region="eu-west-2"}'
        values: "_x60 0+0x180 0+10x120"
    alert_rule_test:
      - alertname: MimirToGrafanaCloudExporterFailures
        eval_time: 70m
      - alertname: MimirToGrafanaCloudExporterFailures
        eval_time: 160m
        exp_alerts:
          - exp_labels:
              area: platform
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "golem"
              customer: "giantswarm"
              installation: "golem"
              namespace: "mimir"
              job: "mimir/mimir-to-grafana-cloud"
              pipeline: "testing"
              provider: "capa"
              region: "eu-west-2"
            exp_annotations:
              dashboard: "promRW001/prometheus-remote-write"
              description: "Prometheus Mimir to Grafana-Cloud is failing to read or write data."
              opsrecipe: "mimir-grafana-cloud-exporter-failing/"
      - alertname: MimirToGrafanaCloudExporterFailures
        eval_time: 200m
      - alertname: MimirToGrafanaCloudExporterFailures
        eval_time: 280m
        exp_alerts:
          - exp_labels:
              area: platform
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "golem"
              customer: "giantswarm"
              installation: "golem"
              namespace: "mimir"
              job: "mimir/mimir-to-grafana-cloud"
              pipeline: "testing"
              provider: "capa"
              region: "eu-west-2"
            exp_annotations:
              dashboard: "promRW001/prometheus-remote-write"
              description: "Prometheus Mimir to Grafana-Cloud is failing to read or write data."
              opsrecipe: "mimir-grafana-cloud-exporter-failing/"
  # Tests for `MimirToGrafanaCloudExporter` alert
  - interval: 1m
    input_series:
      # remote read is working for 2 hours and then fails for 1 hour
      - series: 'kube_pod_status_ready{condition="true", pod="prometheus-mimir-to-grafana-cloud-0", cluster_id="golem", customer="giantswarm", installation="golem", namespace="mimir", pipeline="testing", provider="capa", region="eu-west-2"}'
        values: "_x60 1+0x60 0+0x2 1+0x2 0+0x2 1+0x2 0+0x2 1+x60"
    alert_rule_test:
      - alertname: MimirToGrafanaCloudExporter
        eval_time: 70m
      - alertname: MimirToGrafanaCloudExporter
        eval_time: 140m
        exp_alerts:
          - exp_labels:
              area: platform
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "golem"
              customer: "giantswarm"
              installation: "golem"
              namespace: "mimir"
              pipeline: "testing"
              provider: "capa"
              region: "eu-west-2"
            exp_annotations:
              dashboard: "promRW001/prometheus-remote-write"
              description: "Prometheus Mimir to Grafana-Cloud is restarting too much."
              opsrecipe: "mimir-grafana-cloud-exporter-failing/"
      - alertname: MimirToGrafanaCloudExporterFailures
        eval_time: 180m
