---
rule_files:
  - app.rules.yml

tests:
  # WorkloadClusterAppFailed tests
  - interval: 1m
    input_series:
      - series: 'app_operator_app_info{app="cilium", app_version="1.11.2", catalog="giantswarm", cluster_id="gauss", cluster_missing="false", cluster_type="management_cluster", customer="giantswarm", deployed_version="0.2.6", installation="gauss", instance="100.64.2.221:8000", job="gauss-prometheus/app-exporter-gauss/0", name="cilium", namespace="giantswarm", node="ip-10-0-5-93.eu-west-1.compute.internal", organization="giantswarm", pod="app-exporter-7f8c5d7dd5-9dh4r", provider="aws", service_priority="highest", status="pending-upgrade", team="rocket", upgrade_available="false", version="0.2.6", version_mismatch="false"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: WorkloadClusterAppFailed
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              alertname: WorkloadClusterAppFailed
              app: cilium
              app_version: 1.11.2
              area: managedservices
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              catalog: giantswarm
              cluster_id: gauss
              cluster_missing: "false"
              cluster_type: management_cluster
              customer: giantswarm
              deployed_version: 0.2.6
              installation: gauss
              instance: 100.64.2.221:8000
              job: gauss-prometheus/app-exporter-gauss/0
              name: cilium
              namespace: giantswarm
              node: ip-10-0-5-93.eu-west-1.compute.internal
              organization: giantswarm
              pod: app-exporter-7f8c5d7dd5-9dh4r
              provider: aws
              service_priority: highest
              severity: page
              sig: none
              status: pending-upgrade
              team: rocket
              topic: releng
              upgrade_available: "false"
              version: 0.2.6
              version_mismatch: "false"
            exp_annotations:
              description: "Workload Cluster App giantswarm/cilium, version 0.2.6 is  in pending-upgrade state. "
              opsrecipe: app-failed/

  # AppWithoutTeamAnnotation tests
  - interval: 1m
    input_series:
      - series: 'app_operator_app_info{app="cilium", app_version="1.11.2", catalog="giantswarm", cluster_id="gauss", cluster_missing="false", cluster_type="management_cluster", customer="giantswarm", deployed_version="0.2.6", installation="gauss", instance="100.64.2.221:8000", job="gauss-prometheus/app-exporter-gauss/0", name="cilium", namespace="giantswarm", node="ip-10-0-5-93.eu-west-1.compute.internal", organization="giantswarm", pod="app-exporter-7f8c5d7dd5-9dh4r", provider="aws", service_priority="highest", status="pending-upgrade", team="rocket", upgrade_available="false", version="0.2.6", version_mismatch="false"}'
        values: "1+0x100"
      - series: 'app_operator_app_info{app="userd", catalog="control-plane-catalog", cluster_id="gauss", cluster_missing="false", cluster_type="management_cluster", customer="vodafone", deployed_version="1.2.1", endpoint="web", installation="gauss", instance="app-exporter", job="app-exporter", name="userd", namespace="giantswarm", node="ip-10-0-5-93.eu-west-1.compute.internal", organization="giantswarm", pod="app-exporter-64f9c4fb7b-d77pv", provider="aws", service="app-exporter", service_priority="highest", status="deployed", team="noteam", upgrade_available="false", version="1.2.1", version_mismatch="false"}'
        values: "1+0x100"
    alert_rule_test:
      - alertname: AppWithoutTeamAnnotation
        eval_time: 30m
        exp_alerts:
      - alertname: AppWithoutTeamAnnotation
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              alertname: AppWithoutTeamAnnotation
              app: userd
              area: managedservices
              cancel_if_apiserver_down: true
              cancel_if_cluster_status_creating: true
              cancel_if_cluster_status_deleting: true
              cancel_if_cluster_status_updating: true
              cancel_if_outside_working_hours: true
              catalog: control-plane-catalog
              cluster_id: gauss
              cluster_missing: false
              cluster_type: management_cluster
              customer: vodafone
              deployed_version: 1.2.1
              endpoint: web
              installation: gauss
              instance: app-exporter
              job: app-exporter
              name: userd
              namespace: giantswarm
              node: ip-10-0-5-93.eu-west-1.compute.internal
              organization: giantswarm
              pod: app-exporter-64f9c4fb7b-d77pv
              provider: aws
              service: app-exporter
              service_priority: highest
              severity: notify
              status: deployed
              team: atlas
              topic: releng
              upgrade_available: false
              version: 1.2.1
              version_mismatch: false
            exp_annotations:
              description: App userd has no team label.
              opsrecipe: app-without-team-annotation/

  # ClusterCapiAppFailed tests
  - interval: 1m
    input_series:
      - series: 'app_operator_app_info{app="cluster-azure", catalog="cluster", cluster_missing="true", deployed_version="0.0.15", endpoint="web", instance="192.168.11.76:8000", job="app-exporter", latest_version="0.0.19", name="glippy", namespace="org-giantswarm", pod="app-exporter-9c644b64f-fl7xt", service="app-exporter", status="failed", team="clippy", upgrade_available="true", version="0.0.15", version_mismatch="false"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: ClusterCapiAppFailed
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              alertname: ClusterCapiAppFailed
              app: cluster-azure
              area: kaas
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              catalog: cluster
              cluster_missing: "true"
              deployed_version: 0.0.15
              endpoint: web
              instance: 192.168.11.76:8000
              job: app-exporter
              latest_version: 0.0.19
              name: glippy
              namespace: org-giantswarm
              pod: app-exporter-9c644b64f-fl7xt
              service: app-exporter
              severity: page
              sig: none
              status: failed
              team: clippy
              topic: releng
              upgrade_available: "true"
              version: 0.0.15
              version_mismatch: "false"
            exp_annotations:
              description: "Cluster CAPI App org-giantswarm/glippy, version 0.0.15 is in failed state."
              opsrecipe: app-failed/


