---
rule_files:
  - crsync.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_available{cluster_type="workload_cluster", installation="gazelle", cluster_id="operations", namespace="crsync", deployment="crsync-giantswarm-azurecr-io"}'
        values: "1x5 0x9 1x5 0x10"
    alert_rule_test:
      - alertname: CrsyncDeploymentNotSatisfied
        eval_time: 32m
        exp_alerts:
          - exp_labels:
              alertname: "CrsyncDeploymentNotSatisfied"
              area: "kaas"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "operations"
              cluster_type: "workload_cluster"
              deployment: "crsync-giantswarm-azurecr-io"
              installation: "gazelle"
              namespace: "crsync"
              severity: "page"
              team: "honeybadger"
              topic: "releng"
            exp_annotations:
              description: "CrSync deployment crsync-giantswarm-azurecr-io is not satisfied in gazelle / operations at the crsync namespace."
              opsrecipe: "deployment-not-satisfied/"
