---
rule_files:
  - grafana.management-cluster.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'up{service="grafana", cluster_id="gauss", installation="gauss", instance="grafana"}'
        values: "1+0x20 0+0x100"
    alert_rule_test:
      - alertname: GrafanaDown
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              area: managedservices
              cancel_if_apiserver_down: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_scrape_timeout: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "gauss"
              installation: "gauss"
              instance: "grafana"
              service: "grafana"
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Grafana (grafana) is down."
              opsrecipe: "grafana-down/"
              dashboard: "qRQXmRnik/Grafana"
