rule_files:
  - harbor.all.rules.yml
tests:
  - interval: 1m
    input_series:
      - series: 'harbor_up{instance="harbor", component="trivy"}'
        values: "0+0x10"
    alert_rule_test:
      - alertname: HarborDown
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              app: harbor
              cancel_if_apiserver_down: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_scrape_timeout: "true"
              cancel_if_outside_working_hours: "true"
              instance: harbor
              component: trivy
            exp_annotations:
              description: "harbor (components) trivy has been down for more than 5 minutes."
              summary: " Instance harbor down."
  - interval: 1m
    input_series:
      - series: 'harbor_health{instance="harbor"}'
        values: "0+0x10"
    alert_rule_test:
      - alertname: HarborUnhealthy
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              app: harbor
              cancel_if_apiserver_down: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_scrape_timeout: "true"
              cancel_if_outside_working_hours: "true"
              instance: harbor
            exp_annotations:
              description: "harbor is unhealthy."
              summary: "Habor cluster is unhealthy."
  - interval: 1m
    input_series:
      - series: harbor_project_quota_usage_byte
        values: "0.9+0x10"
      - series: harbor_project_quota_byte
        values: "1+0x10"
    promql_expr_test:
    - expr: (harbor_project_quota_usage_byte) / (harbor_project_quota_byte) > 0.8
      eval_time: 10m
      exp_samples:
        - value: 0.9
  - interval: 1m
    input_series:
      - series: 'kube_pod_status_phase{pod="harbor-exporter", phase="Unknown"}'
        values: "0+0x10"
    alert_rule_test:
      - alertname: harborPodPhaseStatus
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              severity: critical
              app: harbor
              cancel_if_apiserver_down: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_scrape_timeout: "true"
              cancel_if_outside_working_hours: "true"
              pod: "harbor-exporter"
              phase: "Unknown"
            exp_annotations:
              description: "harbor-exporter is an Unknown phase."
              summary: "Pods are not in a succeeded phase."
  
