---
rule_files:
  - inhibit.prometheus-agent.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'aggregation:giantswarm:cluster_release_version{release_version="18.0.2",cluster_id="test1"}'
        values: "1+0x20 0+0x100"
      - series: 'app_operator_app_info{app="observability-bundle",status="deployed",name="test1-observability-bundle",cluster_id="gauss",version="1.1.0"}'
        values: "1+0x20 0+0x100"
    alert_rule_test:
      - alertname: InhibitionVintageClusterReleaseBeforeV19
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              area: empowerment
              team: atlas
              topic: monitoring
              vintage_cluster_release_before_v19: "true"
              release_version: "18.0.2"
              cluster_id: "test1"
            exp_annotations:
              description: "Cluster (test1) using Release (18.0.2)."
      - alertname: InhibitionCAPIClusterIsNotRunningPrometheusAgent
        eval_time: 1m
        exp_alerts:
          - exp_labels:
              area: empowerment
              team: atlas
              topic: monitoring
              capi_cluster_not_running_prometheus_agent: "true"
              app: "observability-bundle"
              status: "deployed"
              name: "test1-observability-bundle"
              cluster_id: "test1"
              version: "1.1.0"
            exp_annotations:
              description: "Cluster (test1) using (observability-bundle) (1.1.0)."
