---
rule_files:
  - prometheus-availability.rules.yml

# Setting evaluation interval to 1h
# to make it faster on long test duration.
evaluation_interval: 1h

tests:
  # Test PrometheusAvailabilityRatio
  - interval: 1m
    input_series:
      # This prometheus is up foreve - generates no alert
      - series: 'kube_pod_status_ready{app="kube-state-metrics", condition="true", container="kube-state-metrics", namespace="install-prometheus", pod="prometheus-install-0"}'
        values: "1+0x120"
      # This prometheus starts at h+1, and takes 5min to get ready - generates no alert
      - series: 'kube_pod_status_ready{app="kube-state-metrics", condition="true", container="kube-state-metrics", namespace="wcok-prometheus", pod="prometheus-wcok-0"}'
        values: "_x60 0+0x5 1+0x60"
      # This prometheus is down - generates alerts
      - series: 'kube_pod_status_ready{app="kube-state-metrics", condition="true", container="kube-state-metrics", namespace="wcbad-prometheus", pod="prometheus-wcbad-0"}'
        values: "0+0x60 1+0x60"
    alert_rule_test:
      - alertname: PrometheusAvailabilityRatio
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_any_apiserver_down: "true"
              cancel_if_cluster_has_no_workers: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              pod: "prometheus-wcbad-0"
            exp_annotations:
              description: "Prometheus prometheus-wcbad-0 has availability ratio of 0.00 (min 0.8) over the last hour."
              opsrecipe: "prometheus-resource-limit-reached/"
              dashboard: "promavailability/prometheus-availability"
      - alertname: PrometheusAvailabilityRatio
        eval_time: 108m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_any_apiserver_down: "true"
              cancel_if_cluster_has_no_workers: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              pod: "prometheus-wcbad-0"
            exp_annotations:
              description: "Prometheus prometheus-wcbad-0 has availability ratio of 0.00 (min 0.8) over the last hour."
              opsrecipe: "prometheus-resource-limit-reached/"
              dashboard: "promavailability/prometheus-availability"
      - alertname: PrometheusAvailabilityRatio
        eval_time: 140m
