---
rule_files:
  - promtail.rules.yml

tests:
  - interval: 1m
    input_series:
      # For the first 60min: test with 1 pod: none, up, down
      - series: 'up{container="promtail",cluster_type="management_cluster", cluster_id="gauss", installation="gauss", node="ip-10-0-5-35.eu-west-1.compute.internal"}'
        values: "_x20 1+0x20 0+0x20" 
      # From 60min: test with 2 pods: 1 up and 1 down, 2 up, 2 down.
      - series: 'up{container="promtail",cluster_type="management_cluster", cluster_id="gauss", installation="gauss", node="ip-10-0-5-145.eu-west-1.compute.internal"}'
        values: "_x60 1+0x20 1+0x20 0+0x20"
      - series: 'up{container="promtail",cluster_type="management_cluster", cluster_id="gauss", installation="gauss", node="ip-10-0-5-76.eu-west-1.compute.internal"}'
        values: "_x60 0+0x20 1+0x20 0+0x20"
    alert_rule_test:
      - alertname: PromtailDown
        eval_time: 10m
      - alertname: PromtailDown
        eval_time: 30m
      - alertname: PromtailDown
        eval_time: 50m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 5 minutes."
              opsrecipe: "promtail-is-not-running/"
      # Tests with 2 pods
      - alertname: PromtailDown
        eval_time: 70m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 5 minutes."
              opsrecipe: "promtail-is-not-running/"
      - alertname: PromtailDown
        eval_time: 90m
      - alertname: PromtailDown
        eval_time: 110m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 5 minutes."
              opsrecipe: "promtail-is-not-running/"
  - interval: 1m
    input_series:
      # Tests with multiple cases: no metrics, no requests, only status_code 204 ones, 204 ones and 500 that are less than 10% of the the total, 500 request that represent more than 10% of the total, only 500 ones
      - series: 'promtail_request_duration_seconds_count{status_code="500", cluster_type="management_cluster", cluster_id="gauss", installation="gauss", node="ip-10-0-5-145.eu-west-1.compute.internal", pod="promtail-2j7z7"}'
        values: "_x60 0+0x60 0+0x60   0+50x60      3000+100x60  9000+600x60"
      - series: 'promtail_request_duration_seconds_count{status_code="204", cluster_type="management_cluster", cluster_id="gauss", installation="gauss", node="ip-10-0-5-145.eu-west-1.compute.internal", pod="promtail-2j7z7"}'
        values: "_x60 0+0x60 0+600x60 36000+600x60 72000+600x60 108000+0x60"
    alert_rule_test:
      - alertname: PromtailRequestsErrors
        eval_time: 30m
      - alertname: PromtailRequestsErrors
        eval_time: 90m
      - alertname: PromtailRequestsErrors
        eval_time: 150m
      - alertname: PromtailRequestsErrors
        eval_time: 210m
      - alertname: PromtailRequestsErrors
        eval_time: 270m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
            exp_annotations:
              description: "This alert checks if that the amount of failed requests is below 10% for promtail"
              opsrecipe: "promtail-requests-are-failing/"
      - alertname: PromtailRequestsErrors
        eval_time: 330m
        exp_alerts:
          - exp_labels:
              area: empowerment
              severity: page
              team: atlas
              topic: observability
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
            exp_annotations:
              description: "This alert checks if that the amount of failed requests is below 10% for promtail"
              opsrecipe: "promtail-requests-are-failing/"
