---
rule_files:
  - sloth.rules.yml

tests:
  - interval: 1m
    input_series:
      # For the first 60min: test with 1 pod: none, up, down
      - series: 'up{app="sloth",cluster_type="management_cluster", cluster_id="gauss", installation="gauss"}'
        values: "_x20 1+0x20 0+0x20"
    alert_rule_test:
      - alertname: SlothDown
        eval_time: 10m
      - alertname: SlothDown
        eval_time: 30m
      - alertname: SlothDown
        eval_time: 50m
        exp_alerts:
          - exp_labels:
              area: managedservices
              cluster_id: gauss
              severity: page
              team: atlas
              topic: observability
              cancel_if_apiserver_down: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_scrape_timeout: "true"
              cancel_if_outside_working_hours: "false"
            exp_annotations:
              description: "Sloth is down."
