---
rule_files:
  - vpa-recommender.rules.yml

tests:
  - interval: 1m
    input_series:
      # For the first 60min: test with 1 pod: none, up, down then pod restarts 3 times in 12min
      - series: 'up{container="recommender"}'
        values: "_x20 1+0x20 0+0x20 1+0x2 0+0x2 1+0x2 0+0x2 1+0x2 0+0x2 1+0x30"
    alert_rule_test:
      - alertname: VpaRecommenderTooManyRestarts
        eval_time: 10m
      - alertname: VpaRecommenderTooManyRestarts
        eval_time: 30m
      - alertname: VpaRecommenderTooManyRestarts
        eval_time: 60m
      - alertname: VpaRecommenderTooManyRestarts
        eval_time: 80m  
      - alertname: VpaRecommenderTooManyRestarts
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              area: managedservices
              severity: page
              team: atlas
              topic: observability
              cancel_if_apiserver_down: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_scrape_timeout: "true"
              cancel_if_outside_working_hours: "false"
            exp_annotations:
              description: "The vpa's recommender has restarted several times in the last 30min"
