---
rule_files:
  - grafana.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'up{service="grafana", cluster_id="golem", cluster_type="management_cluster", installation="golem", instance="grafana"}'
        values: "1+0x20 0+0x100"
    alert_rule_test:
      - alertname: GrafanaDown
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: "golem"
              cluster_type: management_cluster
              installation: "golem"
              instance: "grafana"
              service: "grafana"
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Grafana (grafana) is down."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/grafana-down/
              __dashboardUid__: qRQXmRnik
              dashboardQueryParams: "orgId=2"
  - interval: 1m
    input_series:
      - series: 'cnpg_pg_replication_lag{namespace="monitoring", pod="grafana-postgresql-1", cluster_id="golem", cluster_type="management_cluster", installation="golem"}'
        values: "0+0x20 1+10x30 301+0x20 0+0x20"
    alert_rule_test:
      - alertname: GrafanaPostgresqlReplicationFailure
        eval_time: 15m
        exp_alerts: []
      - alertname: GrafanaPostgresqlReplicationFailure
        eval_time: 30m
        exp_alerts: []
      - alertname: GrafanaPostgresqlReplicationFailure
        eval_time: 65m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              installation: "golem"
              pod: grafana-postgresql-1
              severity: page
              team: atlas
              topic: observability
              cluster_id: "golem"
            exp_annotations:
              description: 'grafana-postgresql replication is lagging for pod grafana-postgresql-1.'
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/grafana-down/
      - alertname: GrafanaPostgresqlReplicationFailure
        eval_time: 85m
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'cnpg_pg_stat_archiver_last_failed_time{namespace="monitoring", pod="grafana-postgresql-1", cluster_id="golem", cluster_type="management_cluster", installation="golem"}'
        values: "-1+0x20 1000+0x10 -1+0x20"
      - series: 'cnpg_pg_stat_archiver_last_archived_time{namespace="monitoring", pod="grafana-postgresql-1", cluster_id="golem", cluster_type="management_cluster", installation="golem"}'
        values: "1000+0x20 -1+0x10 1000+0x20"
    alert_rule_test:
      - alertname: GrafanaPostgresqlArchivingFailure
        eval_time: 15m
        exp_alerts: []
      - alertname: GrafanaPostgresqlArchivingFailure
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              installation: "golem"
              pod: grafana-postgresql-1
              severity: page
              team: atlas
              topic: observability
              cluster_id: "golem"
            exp_annotations:
              description: 'grafana-postgresql archiving failed for pod grafana-postgresql-1.'
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/grafana-down/
      - alertname: GrafanaPostgresqlArchivingFailure
        eval_time: 40m
        exp_alerts: []
