---
rule_files:
  - promtail.rules.yml

tests:
  - interval: 1m
    input_series:
      # For the first 60min: test with 1 pod: none, up, down
      - series: 'up{container="promtail", cluster_id="glean", cluster_type="management_cluster", installation="glean", namespace="kube-system", pod="promtail-1xxxx", provider="capz", pipeline="testing"}'
        values: "_x20 1+0x20 0+0x40"
      - series: kube_pod_info{cluster_id="glean", cluster_type="management_cluster", installation="glean", namespace="kube-system", pod="promtail-1xxxx", node="ip-10-0-5-1.eu-west-1.compute.internal", provider="capz", pipeline="testing"}
        values: "1x180"
      # From 60min: test with 2 pods: 1 up and 1 down, 2 up, 2 down.
      - series: 'up{container="promtail", cluster_id="glean", cluster_type="management_cluster", installation="glean", namespace="kube-system", pod="promtail-2xxxx", provider="capz", pipeline="testing"}'
        values: "_x80 1+0x40 1+0x20 0+0x40"
      - series: kube_pod_info{cluster_id="glean", cluster_type="management_cluster", installation="glean", namespace="kube-system", pod="promtail-2xxxx", node="ip-10-0-5-2.eu-west-1.compute.internal", provider="capz", pipeline="testing"}
        values: "1x180"
      - series: 'up{container="promtail", cluster_type="management_cluster", cluster_id="glean", installation="glean", namespace="kube-system", pod="promtail-3xxxx", provider="capz", pipeline="testing"}'
        values: "_x80 0+0x40 1+0x20 0+0x40"
      - series: kube_pod_info{cluster_id="glean", cluster_type="management_cluster", installation="glean", namespace="kube-system", pod="promtail-3xxxx", node="ip-10-0-5-3.eu-west-1.compute.internal", provider="capz", pipeline="testing"}
        values: "1x180"
    alert_rule_test:
      - alertname: PromtailDown
        eval_time: 10m
      - alertname: PromtailDown
        eval_time: 30m
      - alertname: PromtailDown
        eval_time: 71m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: glean
              cluster_type: management_cluster
              installation: glean
              namespace: kube-system
              node: ip-10-0-5-1.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-1xxxx
              provider: capz
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      # Tests with 2 pods
      - alertname: PromtailDown
        eval_time: 111m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: glean
              cluster_type: management_cluster
              installation: glean
              namespace: kube-system
              node: ip-10-0-5-3.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-3xxxx
              provider: capz
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      - alertname: PromtailDown
        eval_time: 121m
      - alertname: PromtailDown
        eval_time: 180m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: glean
              cluster_type: management_cluster
              installation: glean
              namespace: kube-system
              node: ip-10-0-5-2.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-2xxxx
              provider: capz
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: glean
              cluster_type: management_cluster
              installation: glean
              namespace: kube-system
              node: ip-10-0-5-3.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-3xxxx
              provider: capz
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
  - interval: 1m
    input_series:
      # Tests with multiple cases: no metrics, no requests, only status_code 204 ones, 204 ones and 500 that are less than 10% of the the total, 500 request that represent more than 10% of the total, only 500 ones
      - series: 'promtail_request_duration_seconds_count{status_code="500", cluster_type="management_cluster", cluster_id="glean", installation="glean", provider="capz", pipeline="testing", node="ip-10-0-5-145.eu-west-1.compute.internal", pod="promtail-2j7z7"}'
        values: "_x60 0+0x60 0+0x60   0+50x60      3000+100x60  9000+600x60"
      - series: 'promtail_request_duration_seconds_count{status_code="204", cluster_type="management_cluster", cluster_id="glean", installation="glean", provider="capz", pipeline="testing", node="ip-10-0-5-145.eu-west-1.compute.internal", pod="promtail-2j7z7"}'
        values: "_x60 0+0x60 0+600x60 36000+600x60 72000+600x60 108000+0x60"
    alert_rule_test:
      - alertname: PromtailRequestsErrors
        eval_time: 30m
      - alertname: PromtailRequestsErrors
        eval_time: 90m
      - alertname: PromtailRequestsErrors
        eval_time: 150m
      - alertname: PromtailRequestsErrors
        eval_time: 210m
      - alertname: PromtailRequestsErrors
        eval_time: 270m
      - alertname: PromtailRequestsErrors
        eval_time: 310m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cluster_id: glean
              installation: glean
              provider: capz
              pipeline: testing
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              __dashboardUid__: promtail-overview
              dashboardQueryParams: "orgId=2"
              description: "This alert checks if that the amount of failed requests is below 10% for promtail"
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      - alertname: PromtailRequestsErrors
        eval_time: 330m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cluster_id: glean
              installation: glean
              provider: capz
              pipeline: testing
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              __dashboardUid__: promtail-overview
              dashboardQueryParams: "orgId=2"
              description: "This alert checks if that the amount of failed requests is below 10% for promtail"
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
