---
rule_files:
  - crossplane.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_unavailable{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="crossplane", installation="gauss", instance="100.64.5.122:8080", job="gauss-prometheus/workload-gauss/0", namespace="crossplane", node="ip-10-0-5-119.eu-west-1.compute.internal", organization="giantswarm", pod="kube-state-metrics-95bbb4bd7-v6hvh", provider="aws", service_priority="highest"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: CrossplaneDeploymentNotSatisfied
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              alertname: CrossplaneDeploymentNotSatisfied
              app: kube-state-metrics
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: gauss
              cluster_type: management_cluster
              container: kube-state-metrics
              customer: giantswarm
              deployment: crossplane
              installation: gauss
              instance: 100.64.5.122:8080
              job: gauss-prometheus/workload-gauss/0
              namespace: crossplane
              node: ip-10-0-5-119.eu-west-1.compute.internal
              organization: giantswarm
              pod: kube-state-metrics-95bbb4bd7-v6hvh
              provider: aws
              service_priority: highest
              severity: page
              team: honeybadger
              topic: managementcluster
            exp_annotations:
              description: "Crossplane related deployment crossplane/crossplane is not satisfied."
              opsrecipe: "deployment-not-satisfied/"
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_unavailable{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="caicloud-event-exporter", installation="gauss", instance="100.64.5.122:8080", job="gauss-prometheus/workload-gauss/0", namespace="crossplane", node="ip-10-0-5-119.eu-west-1.compute.internal", organization="giantswarm", pod="kube-state-metrics-95bbb4bd7-v6hvh", provider="aws", service_priority="highest"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: CrossplaneDeploymentNotSatisfied
        eval_time: 51m
        exp_alerts:
          - exp_labels:
              alertname: CrossplaneDeploymentNotSatisfied
              app: kube-state-metrics
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: gauss
              cluster_type: management_cluster
              container: kube-state-metrics
              customer: giantswarm
              deployment: caicloud-event-exporter
              installation: gauss
              instance: 100.64.5.122:8080
              job: gauss-prometheus/workload-gauss/0
              namespace: crossplane
              node: ip-10-0-5-119.eu-west-1.compute.internal
              organization: giantswarm
              pod: kube-state-metrics-95bbb4bd7-v6hvh
              provider: aws
              service_priority: highest
              severity: page
              team: honeybadger
              topic: managementcluster
            exp_annotations:
              description: "Crossplane related deployment crossplane/caicloud-event-exporter is not satisfied."
              opsrecipe: "deployment-not-satisfied/"
