---
rule_files:
  - external-secrets.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_unavailable{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="external-secrets", installation="gauss", instance="100.64.6.226:8080", job="gauss-prometheus/workload-gauss/0", namespace="external-secrets", node="ip-10-0-5-161.eu-west-1.compute.internal", organization="giantswarm", pod="kube-state-metrics-fd99568b6-fnhdv", provider="aws", service_priority="highest"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: DeploymentNotSatisfiedExternalSecrets
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              alertname: DeploymentNotSatisfiedExternalSecrets
              app: kube-state-metrics
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: gauss
              cluster_type: management_cluster
              container: kube-state-metrics
              customer: giantswarm
              deployment: external-secrets
              installation: gauss
              instance: 100.64.6.226:8080
              job: gauss-prometheus/workload-gauss/0
              namespace: external-secrets
              node: ip-10-0-5-161.eu-west-1.compute.internal
              organization: giantswarm
              pod: kube-state-metrics-fd99568b6-fnhdv
              provider: aws
              service_priority: highest
              severity: page
              team: honeybadger
              topic: managementcluster
            exp_annotations:
              description: "ExternalSecrets related deployment external-secrets/external-secrets is not satisfied."
              opsrecipe: "deployment-not-satisfied/"
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_unavailable{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="external-secrets-cert-controller", installation="gauss", instance="100.64.6.226:8080", job="gauss-prometheus/workload-gauss/0", namespace="external-secrets", node="ip-10-0-5-161.eu-west-1.compute.internal", organization="giantswarm", pod="kube-state-metrics-fd99568b6-fnhdv", provider="aws", service_priority="highest"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: DeploymentNotSatisfiedExternalSecrets
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              alertname: DeploymentNotSatisfiedExternalSecrets
              app: kube-state-metrics
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: gauss
              cluster_type: management_cluster
              container: kube-state-metrics
              customer: giantswarm
              deployment: external-secrets-cert-controller
              installation: gauss
              instance: 100.64.6.226:8080
              job: gauss-prometheus/workload-gauss/0
              namespace: external-secrets
              node: ip-10-0-5-161.eu-west-1.compute.internal
              organization: giantswarm
              pod: kube-state-metrics-fd99568b6-fnhdv
              provider: aws
              service_priority: highest
              severity: page
              team: honeybadger
              topic: managementcluster
            exp_annotations:
              description: "ExternalSecrets related deployment external-secrets/external-secrets-cert-controller is not satisfied."
              opsrecipe: "deployment-not-satisfied/"
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_unavailable{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="external-secrets-webhook", installation="gauss", instance="100.64.6.226:8080", job="gauss-prometheus/workload-gauss/0", namespace="external-secrets", node="ip-10-0-5-161.eu-west-1.compute.internal", organization="giantswarm", pod="kube-state-metrics-fd99568b6-fnhdv", provider="aws", service_priority="highest"}'
        values: "0+0x20 1+0x100"
    alert_rule_test:
      - alertname: DeploymentNotSatisfiedExternalSecrets
        eval_time: 60m
        exp_alerts:
          - exp_labels:
              alertname: DeploymentNotSatisfiedExternalSecrets
              app: kube-state-metrics
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: gauss
              cluster_type: management_cluster
              container: kube-state-metrics
              customer: giantswarm
              deployment: external-secrets-webhook
              installation: gauss
              instance: 100.64.6.226:8080
              job: gauss-prometheus/workload-gauss/0
              namespace: external-secrets
              node: ip-10-0-5-161.eu-west-1.compute.internal
              organization: giantswarm
              pod: kube-state-metrics-fd99568b6-fnhdv
              provider: aws
              service_priority: highest
              severity: page
              team: honeybadger
              topic: managementcluster
            exp_annotations:
              description: "ExternalSecrets related deployment external-secrets/external-secrets-webhook is not satisfied."
              opsrecipe: "deployment-not-satisfied/"
