---
rule_files:
  - flux.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'gotk_resource_info{ready="False", job="giantswarm/cluster-api-monitoring", customresource_kind="Kustomization", cluster_type="management_cluster", exported_namespace="flux-giantswarm", name="silences"}'
        values: "1x60"
    alert_rule_test:
      - alertname: FluxKustomizationFailed
        eval_time: 30m
        exp_alerts: []
  - interval: 1m
    input_series:
      - series: 'gotk_resource_info{installation="test", job="giantswarm/cluster-api-monitoring", cluster_type="management_cluster", exported_namespace="flux-giantswarm", customresource_kind="Kustomization", name="flux", suspended="true"}'
        values: "1x60 0+1x60 1+0x1500"
    alert_rule_test:
      - alertname:  FluxSuspendedForTooLong
        eval_time: 1560m
        exp_alerts:
          - exp_labels:
              alertname: "FluxSuspendedForTooLong"
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_type: "management_cluster"
              customresource_kind: "Kustomization"
              exported_namespace: "flux-giantswarm"
              installation: "test"
              job: "giantswarm/cluster-api-monitoring"
              name: "flux"
              severity: "page"
              suspended: "true"
              team: "honeybadger"
              topic: "releng"
            exp_annotations:
              description: "Flux Kustomization flux in ns flux-giantswarm on test has been suspended for 24h."
              runbook_url: "https://intranet.giantswarm.io/docs/support-and-ops/runbooks/flux-suspended/?INSTALLATION=test&CLUSTER="
