---
rule_files:
  - zot.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'kube_deployment_status_replicas_unavailable{cluster_type="management_cluster",namespace="zot",deployment="zot-zot"}'
        values: '_x5 0+0x10 1+0x45'
    alert_rule_test:
      - alertname: DeploymentNotSatisfiedZot
        eval_time: 46m
        exp_alerts:
          - exp_labels:
              alertname: "DeploymentNotSatisfiedZot"
              area: "platform"
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "false"
              cluster_type: "management_cluster"
              deployment: "zot-zot"
              namespace: "zot"
              severity: "page"
              team: "honeybadger"
              topic: "managementcluster"
            exp_annotations:
              description: "Zot deployment zot/zot-zot is not satisfied."
              opsrecipe: "zot/"
