---
rule_files:
  - kong.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'up{job="kong_datastore_reachable", namespace="giantswarm", provider="aws"}'
        values: '0+0x10'
    alert_rule_test:
      - alertname: KongDatastoreNotReachable
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              alertname: KongDatastoreNotReachable
              area: managedservices
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_outside_working_hours: "false"
              job: kong_datastore_reachable
              namespace: giantswarm
              provider: aws
              severity: page
              team: cabbage
              topic: kong
            exp_annotations:
              description: "Datastore of Kong in namespace giantswarm is not reachable."
              opsrecipe: managed-app-kong/#kong-datastore-not-reachable
