---
rule_files:
  - statefulset.management-cluster.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'kube_statefulset_status_replicas{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", customer="giantswarm", installation="gauss", namespace="loki", organization="giantswarm", pipeline="testing", region="westeurope", statefulset="loki-write"}'
        values: "3+0x180"
      - series: 'kube_statefulset_status_replicas_ready{app="kube-state-metrics", cluster_id="gauss", cluster_type="management_cluster", customer="giantswarm", installation="gauss", namespace="loki", organization="giantswarm", pipeline="testing", region="westeurope", statefulset="loki-write"}'
        values: "3+0x60 2+0x60 3+0x60"
    alert_rule_test:
      - alertname: StatefulsetNotSatisfiedAtlas
        eval_time: 30m
      - alertname: StatefulsetNotSatisfiedAtlas
        eval_time: 60m
      - alertname: StatefulsetNotSatisfiedAtlas
        eval_time: 90m
      - alertname: StatefulsetNotSatisfiedAtlas
        eval_time: 100m
        exp_alerts:
          - exp_labels:
              app: kube-state-metrics
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "false"
              cluster_id: "gauss"
              cluster_type: management_cluster
              customer: giantswarm
              installation: "gauss"
              namespace: loki
              organization: giantswarm
              pipeline: "testing"
              region: westeurope
              severity: page
              statefulset: loki-write
              team: atlas
              topic: managementcluster
            exp_annotations:
              description: "Statefulset loki/loki-write is not satisfied."
              opsrecipe: "deployment-not-satisfied/"
      - alertname: StatefulsetNotSatisfiedAtlas
        eval_time: 150m
