rule_files:
  - capo.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_openstackmachine_status_instance_state{cluster="galaxy", name="galaxy-72jq5", exported_namespace="giantswarm", state="ACTIVE"}'
        values: "1+0x10 0+0x35"
      - series: 'capi_openstackmachine_status_instance_state{cluster="galaxy", name="galaxy-72jq5", exported_namespace="giantswarm", state="ERROR"}'
        values: "0+0x10 1+0x35"
      - series: 'capi_openstackmachine_status_instance_state{cluster="galaxy", name="galaxy-9xjq5", exported_namespace="giantswarm", state="ERROR"}'
        values: "1+0x60"
    alert_rule_test:
      - alertname: OpenStackMachineUnexpectedInstanceState
        eval_time: 1h
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: rocket
              topic: managementcluster
              cluster: galaxy
              name: galaxy-9xjq5
              exported_namespace: giantswarm
              state: ERROR
            exp_annotations:
              description: "OpenStackMachine giantswarm/galaxy-9xjq5 instance state is ERROR."
  - interval: 1m
    input_series:
      - series: 'capi_openstackmachine_status_failure{cluster="galaxy", name="galaxy-72jq5", exported_namespace="giantswarm"}'
        values: "0+0x45"
      - series: 'capi_openstackmachine_status_failure{cluster="galaxy", name="galaxy-98jq5", exported_namespace="giantswarm", reason="transient error reason"}'
        values: "1+0x45"
    alert_rule_test:
      - alertname: OpenStackMachineFailure
        eval_time: 45m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: rocket
              topic: managementcluster
              cluster: galaxy
              name: galaxy-98jq5
              exported_namespace: giantswarm
              reason: transient error reason
            exp_annotations:
              description: "OpenStackMachine giantswarm/galaxy-98jq5 failure reason is transient error reason."
              opsrecipe: "remove-errors-from-capi-capo-crs/"
  - interval: 1m
    input_series:
      - series: 'capi_openstackcluster_status_failure{cluster="galaxy", name="galaxy-72jq5", exported_namespace="giantswarm"}'
        values: "0+0x45"
      - series: 'capi_openstackcluster_status_failure{cluster="galaxy", name="galaxy-9xjq5", exported_namespace="giantswarm", reason="LB stuck in PENDING_UPDATE"}'
        values: "1+0x45"
    alert_rule_test:
      - alertname: OpenStackClusterFailure
        eval_time: 45m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: rocket
              topic: managementcluster
              cluster: galaxy
              name: galaxy-9xjq5
              exported_namespace: giantswarm
              reason: LB stuck in PENDING_UPDATE
            exp_annotations:
              description: "OpenStackCluster giantswarm/galaxy-9xjq5 failure reason is LB stuck in PENDING_UPDATE."
              opsrecipe: "remove-errors-from-capi-capo-crs/"
