---
rule_files:
  - kyverno.rules.yml
tests:
  - interval: 1m
    input_series:
      # Kyverno validating webhooks
      - series: 'kube_validatingwebhookconfiguration_info{cluster_id="gremlin", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", endpoint="http", installation="gremlin", instance="10.0.135.241:8080", job="kube-state-metrics", node="master-00000y", organization="giantswarm", pipeline="testing", pod="prometheus-operator-app-kube-state-metrics-d7f4ff68d-qn6sb", prometheus="kube-system/prometheus-agent", prometheus_replica="prometheus-prometheus-agent-0", provider="capa", region="germanywestcentral", service="prometheus-operator-app-kube-state-metrics", service_priority="highest", validatingwebhookconfiguration="kyverno-resource-validating-webhook-cfg"}'
        values: "1+0x20"
      # Kyverno deployment status replicas
      - series: 'kube_deployment_status_replicas_ready{cluster_id="gremlin", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="kyverno", endpoint="http", installation="gremlin", instance="10.0.135.241:8080", job="kube-state-metrics", namespace="kyverno", node="master-00000y", organization="giantswarm", pipeline="testing", pod="prometheus-operator-app-kube-state-metrics-d7f4ff68d-qn6sb", prometheus="kube-system/prometheus-agent", prometheus_replica="prometheus-prometheus-agent-0", provider="capa", region="germanywestcentral", service="prometheus-operator-app-kube-state-metrics", service_priority="highest"}'
        values: "0+0x20"
      # Kyverno deployment spec replicas
      - series: 'kube_deployment_spec_replicas{cluster_id="gremlin", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", deployment="kyverno", endpoint="http", installation="gremlin", instance="10.0.135.241:8080", job="kube-state-metrics", namespace="kyverno", node="master-00000y", organization="giantswarm", pipeline="testing", pod="prometheus-operator-app-kube-state-metrics-d7f4ff68d-qn6sb", prometheus="kube-system/prometheus-agent", prometheus_replica="prometheus-prometheus-agent-0", provider="capa", region="germanywestcentral", service="prometheus-operator-app-kube-state-metrics", service_priority="highest"}'
        values: "0+0x240 1+0x70"
      # Kyverno admission reports
      - series: 'aggregation:kyverno_resource_counts{cluster_id="gremlin", installation="gremlin", kind="admissionreports.kyverno.io", pipeline="testing", provider="capa"}'
        values: "0+1000x30 30000+1500x30"
      # Kyverno updaterequests
      - series: 'aggregation:kyverno_resource_counts{cluster_id="gremlin", installation="gremlin", kind="updaterequests.kyverno.io", pipeline="testing", provider="capa"}'
        values: "0+100x15 5000+1500x30"
    alert_rule_test:
      # Webhooks alert
      - alertname: KyvernoWebhookHasNoAvailableReplicas
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              area: platform
              cluster_id: gremlin
              installation: gremlin
              pipeline: testing
              provider: capa
              severity: page
              team: shield
              topic: kyverno
              cancel_if_outside_working_hours: "false"
            exp_annotations:
              description: "Kyverno has no available replicas but webhooks are present."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/kyverno-webhooks/
      # Kyverno reports too high alert
      - alertname: KyvernoUpdateRequestsCountTooHigh
        eval_time: 45m
        exp_alerts:
          - exp_labels:
              area: platform
              cluster_id: gremlin
              installation: gremlin
              pipeline: testing
              provider: capa
              severity: notify
              team: shield
              topic: kyverno
              cancel_if_outside_working_hours: "false"
              kind: "updaterequests.kyverno.io"
            exp_annotations:
              description: "Kyverno updaterequests.kyverno.io are too high. This is an indicator that Kyverno's background controller may not be able to create some resources."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/kyverno-updaterequests/
      # Kyverno scaled down alert
      - alertname: KyvernoScaledDownTooLong
        eval_time: 240m
        exp_alerts:
          - exp_labels:
              area: platform
              cluster_id: gremlin
              installation: gremlin
              pipeline: testing
              provider: capa
              severity: notify
              team: shield
              topic: kyverno
              cancel_if_outside_working_hours: "true"
            exp_annotations:
              description: "Kyverno has been scaled down for too long."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/kyverno-scaled-down/
      # Kyverno unsafe replica count alert
      - alertname: KyvernoUnsafeReplicaCount
        eval_time: 310m
        exp_alerts:
          - exp_labels:
              area: platform
              cluster_id: gremlin
              installation: gremlin
              pipeline: testing
              provider: capa
              severity: notify
              team: shield
              topic: kyverno
              cancel_if_outside_working_hours: "true"
            exp_annotations:
              description: "Kyverno's admission controller deployment must use at least 3 replicas, or be scaled to 0."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/KyvernoWronglyScaled/
