---
rule_files:
  - 'teleport.rules.yml'

tests:
  - interval: 1m
    input_series:
    - series: 'kube_secret_created{cluster_id="my-cluster", installation="golem", secret="my-cluster-teleport-join-token"}'
      values: "1+0x150"
    - series: 'capi_cluster_status_phase{cluster_id="my-cluster", installation="golem", phase="Provisioned"}'
      values: "2+0x150"
    alert_rule_test:
      - alertname: TeleportJoinTokenSecretMismatch
        eval_time: 30m
        exp_alerts: []
      - alertname: TeleportJoinTokenSecretMismatch
        eval_time: 140m
        exp_alerts:
          - exp_labels:
              alertname: TeleportJoinTokenSecretMismatch
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: my-cluster
              installation: golem
              severity: notify
              team: shield
              topic: teleport
            exp_annotations:
              description: "Mismatch in number of teleport-join-token secrets and clusters"
  - interval: 1m
    input_series:
      - series: 'kube_configmap_info{cluster_id="my-cluster", installation="grizzly", configmap="my-cluster-teleport-kube-agent-config"}'
        values: "1+0x150"
      - series: 'capi_cluster_status_phase{cluster_id="my-cluster", installation="grizzly", phase="Provisioned"}'
        values: "2+0x150"
    alert_rule_test:
      - alertname: TeleportKubeAgentConfigMapMismatch
        eval_time: 30m
        exp_alerts: []
      - alertname: TeleportKubeAgentConfigMapMismatch
        eval_time: 140m
        exp_alerts:
          - exp_labels:
              alertname: TeleportKubeAgentConfigMapMismatch
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: my-cluster
              installation: grizzly
              severity: notify
              team: shield
              topic: teleport
            exp_annotations:
              description: "Mismatch in number of teleport-kube-agent-config secrets and clusters"
  - interval: 1m
    input_series:
      - series: 'kube_statefulset_status_replicas{cluster_id="my-cluster", installation="gunittest", statefulset="teleport-kube-agent"}'
        values: "3+0x60"
      - series: 'kube_statefulset_status_replicas_ready{cluster_id="my-cluster", installation="gunittest", statefulset="teleport-kube-agent"}'
        values: "2+0x60"
    alert_rule_test:
      - alertname: TeleportKubeAgentInstancesNotReady
        eval_time: 30m
        exp_alerts: []
      - alertname: TeleportKubeAgentInstancesNotReady
        eval_time: 50m
        exp_alerts:
          - exp_labels:
              alertname: TeleportKubeAgentInstancesNotReady
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: my-cluster
              installation: gunittest
              severity: notify
              statefulset: teleport-kube-agent
              team: shield
              topic: teleport
            exp_annotations:
              description: "At least one teleport-kube-agent instance is unhealthy."
              summary: "Teleport Kube Agent has had one or more unhealthy instances for an extended period."
              runbook_url: "https://intranet.giantswarm.io/docs/support-and-ops/teleport/"
  - interval: 1m
    input_series:
      - series: 'kube_statefulset_status_replicas{cluster_id="zero-test", installation="gunittest", statefulset="teleport-kube-agent"}'
        values: "3+0x40"
      - series: 'kube_statefulset_status_replicas_ready{cluster_id="zero-test", installation="gunittest", statefulset="teleport-kube-agent"}'
        values: "0+0x40"
    alert_rule_test:
      - alertname: TeleportKubeAgentZeroReadyReplicas
        eval_time: 20m
        exp_alerts: []
      - alertname: TeleportKubeAgentZeroReadyReplicas
        eval_time: 30m
        exp_alerts:
          - exp_labels:
              alertname: TeleportKubeAgentZeroReadyReplicas
              area: platform
              cancel_if_outside_working_hours: "false"
              cluster_id: zero-test
              installation: gunittest
              severity: notify
              statefulset: teleport-kube-agent
              team: shield
              topic: teleport
            exp_annotations:
              description: "Cluster has no Ready teleport-kube-agent instances."
              summary: "Teleport Kube Agent has had no Ready instances for an extended period."
              runbook_url: "https://intranet.giantswarm.io/docs/support-and-ops/teleport/"
