---
rule_files:
  - certificate.management-cluster.rules.yml

tests:
  - interval: 5m
    input_series:
      - series: 'cert_exporter_secret_not_after{secretkey="tls.crt", certificatename="aws-admission-controller-certificates"}'
        values: "1+0x20 0+0x100"
      - series: 'cert_exporter_certificate_cr_not_after{name="aws-admission-controller-certificates"}'
        values: "1+0x20 0+0x100"
    alert_rule_test:
      - alertname: ManagementClusterCertificateIsMissing
        eval_time: 5m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "true"
              severity: page
              team: phoenix
              topic: security
            exp_annotations:
              description: 'Cannot renew Certificate for Secret giantswarm/aws-admission-controller-certificates on gauss because it is missing.'
              opsrecipe: managed-app-cert-manager/missing-certificate-for-secret/
