---
rule_files:
  - promtail.rules.yml

tests:
  - interval: 1m
    input_series:
      # For the first 60min: test with 1 pod: none, up, down
      - series: 'up{container="promtail", cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="promtail-1xxxx", provider="aws", pipeline="testing"}'
        values: "_x20 1+0x20 0+0x40"
      - series: kube_pod_info{cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="promtail-1xxxx", node="ip-10-0-5-1.eu-west-1.compute.internal", provider="aws", pipeline="testing"}
        values: "1x180"
      # From 60min: test with 2 pods: 1 up and 1 down, 2 up, 2 down.
      - series: 'up{container="promtail", cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="promtail-2xxxx", provider="aws", pipeline="testing"}'
        values: "_x80 1+0x40 1+0x20 0+0x40"
      - series: kube_pod_info{cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="promtail-2xxxx", node="ip-10-0-5-2.eu-west-1.compute.internal", provider="aws", pipeline="testing"}
        values: "1x180"
      - series: 'up{container="promtail", cluster_type="management_cluster", cluster_id="golem", installation="golem", namespace="kube-system", pod="promtail-3xxxx", provider="aws", pipeline="testing"}'
        values: "_x80 0+0x40 1+0x20 0+0x40"
      - series: kube_pod_info{cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="promtail-3xxxx", node="ip-10-0-5-3.eu-west-1.compute.internal", provider="aws", pipeline="testing"}
        values: "1x180"
    alert_rule_test:
      - alertname: PromtailDown
        eval_time: 10m
      - alertname: PromtailDown
        eval_time: 30m
      - alertname: PromtailDown
        eval_time: 71m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: golem
              cluster_type: management_cluster
              installation: golem
              namespace: kube-system
              node: ip-10-0-5-1.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-1xxxx
              provider: aws
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      # Tests with 2 pods
      - alertname: PromtailDown
        eval_time: 111m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: golem
              cluster_type: management_cluster
              installation: golem
              namespace: kube-system
              node: ip-10-0-5-3.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-3xxxx
              provider: aws
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      - alertname: PromtailDown
        eval_time: 121m
      - alertname: PromtailDown
        eval_time: 180m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: golem
              cluster_type: management_cluster
              installation: golem
              namespace: kube-system
              node: ip-10-0-5-2.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-2xxxx
              provider: aws
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cancel_if_node_unschedulable: "true"
              cancel_if_node_not_ready: "true"
              cluster_id: golem
              cluster_type: management_cluster
              installation: golem
              namespace: kube-system
              node: ip-10-0-5-3.eu-west-1.compute.internal
              pipeline: testing
              pod: promtail-3xxxx
              provider: aws
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Scraping of all promtail pods to check if one failed every 30 minutes."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
  - interval: 1m
    input_series:
      # Tests with multiple cases:
      #   no metrics (no alert)
      #   no requests (no alert)
      #   only status_code 204 ones (no alert)
      #   status codes 204 and less than 10% of status codes 500 (no alert)
      #   status codes 500 requests that represent more than 10% of the total (alert fires)
      #   only status codes 500 (alert fires)
      #   only status code 500, but alloy-logs is installed (no alert)
      - series: 'promtail_request_duration_seconds_count{status_code="500", cluster_type="management_cluster", cluster_id="golem", installation="golem", provider="aws", pipeline="testing", node="ip-10-0-5-145.eu-west-1.compute.internal", pod="promtail-2j7z7"}'
        values: "_x60 0+0x60 0+0x60   0+50x60      3000+100x60  9000+600x60 45000+600x60"
      - series: 'promtail_request_duration_seconds_count{status_code="204", cluster_type="management_cluster", cluster_id="golem", installation="golem", provider="aws", pipeline="testing", node="ip-10-0-5-145.eu-west-1.compute.internal", pod="promtail-2j7z7"}'
        values: "_x60 0+0x60 0+600x60 36000+600x60 72000+600x60 108000+0x60 108000+0x60"
      - series: 'up{service="alloy-logs", cluster_id="golem"}'
        values: "_x60 _x60   _x60     _x60         _x60         _x60        1+0x60"
    alert_rule_test:
      - alertname: PromtailRequestsErrors
        eval_time: 30m
      - alertname: PromtailRequestsErrors
        eval_time: 90m
      - alertname: PromtailRequestsErrors
        eval_time: 150m
      - alertname: PromtailRequestsErrors
        eval_time: 210m
      - alertname: PromtailRequestsErrors
        eval_time: 270m
      - alertname: PromtailRequestsErrors
        eval_time: 310m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: golem
              installation: golem
              provider: aws
              pipeline: testing
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              __dashboardUid__: promtail-overview
              dashboardQueryParams: "orgId=2"
              description: "This alert checks if that the amount of failed requests is below 10% for promtail"
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      - alertname: PromtailRequestsErrors
        eval_time: 330m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: golem
              installation: golem
              provider: aws
              pipeline: testing
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              __dashboardUid__: promtail-overview
              dashboardQueryParams: "orgId=2"
              description: "This alert checks if that the amount of failed requests is below 10% for promtail"
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/
      - alertname: PromtailRequestsErrors
        eval_time: 360m
  - interval: 1m
    input_series:
      # Test various configs of promtail + alloy-logs:
      #  - 120min with none installed
      #  - 120min with only promtail
      #  - 120min with promtail + alloy
      #  - 120min with only alloy
      - series: 'up{service="promtail-metrics", cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="promtail-1xxxx", provider="capa", pipeline="testing"}'
        values: "_x120 1+0x120 1+0x120 _x120"
      - series: 'up{service="alloy-logs", cluster_id="golem", cluster_type="management_cluster", installation="golem", namespace="kube-system", pod="alloy-logs-1xxxx", provider="capa", pipeline="testing"}'
        values: "_x120 _x120   1+0x120 1+0x120"
    alert_rule_test:
      - alertname: PromtailConflictsWithAlloy
        eval_time: 70m
      - alertname: PromtailConflictsWithAlloy
        eval_time: 190m
      - alertname: PromtailConflictsWithAlloy
        eval_time: 310m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_outside_working_hours: "true"
              cluster_id: golem
              installation: golem
              pipeline: testing
              provider: capa
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              __dashboardUid__: promtail-overview
              dashboardQueryParams: "orgId=2"
              description: "Both promtail and alloy-logs are installed and conflicting on the cluster"
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/promtail/#check-that-it-does-not-conflict-with-alloy-logs
      - alertname: PromtailConflictsWithAlloy
        eval_time: 430m
