rule_files:
  - capi-kubeadmcontrolplane.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_kubeadmcontrolplane_spec_replicas{cluster_name="clippaxy", name="clippaxy-72jq5", exported_namespace="giantswarm"}'
        values: "0+3x100"
      - series: 'capi_kubeadmcontrolplane_status_replicas_ready{cluster_name="clippaxy", name="clippaxy-72jq5", exported_namespace="giantswarm"}'
        values: "0+3x100"
      - series: 'capi_kubeadmcontrolplane_spec_replicas{cluster_name="clippaxy", name="clippaxy-72jzy", exported_namespace="giantswarm"}'
        values: "0+3x100"
      - series: 'capi_kubeadmcontrolplane_status_replicas_ready{cluster_name="clippaxy", name="clippaxy-72jzy", exported_namespace="giantswarm"}'
        values: "0+2x100"
      - series: 'capi_kubeadmcontrolplane_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-72r5c", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: KubeadmControlPlaneReplicasMismatch
        eval_time: 100m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_prometheus_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: clippaxy
              name: clippaxy-72jzy
              exported_namespace: giantswarm
            exp_annotations:
              description: "The clusters clippaxy kubeadmcontrolplane giantswarm/clippaxy-72jzy does not match the expected number of replicas for longer than 90 minutes."
              opsrecipe: capi-kubeadmcontrolplane/
              dashboard: bdi7iswg81czkcasd/capi-agregated-error-logs-for-capi-controllers
      - alertname: KubeadmControlPlanePaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_prometheus_agent_down: "true"
              cancel_if_outside_working_hours: "true"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-72r5c
              exported_namespace: giantswarm
              paused_value: true
            exp_annotations:
              description: "The clusters grumpy kubeadmcontrolplane giantswarm/grumpy-72r5c is paused."
              opsrecipe: capi-kubeadmcontrolplane/
              dashboard: bdi7iswg81czkcasd/capi-agregated-error-logs-for-capi-controllers
