rule_files:
  - capi-machinedeployment.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machinedeployment_spec_replicas{cluster_name="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinedeployment_status_replicas_available{cluster_name="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinedeployment_spec_replicas{cluster_name="clippaxy", name="clippaxy-def99", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinedeployment_status_replicas_available{cluster_name="clippaxy", name="clippaxy-def99", exported_namespace="giantswarm"}'
        values: "0+2x75"
      - series: 'capi_machinedeployment_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-def99", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: MachineDeploymentReplicasMismatch
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "false"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: clippaxy
              name: clippaxy-def99
              exported_namespace: giantswarm
            exp_annotations:
              description: "The clusters clippaxy machinedeployment giantswarm/clippaxy-def99 does not match the expected number of replicas for longer than 1h."
              opsrecipe: capi-machinedeployment/
      - alertname: MachineDeploymentPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "false"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-def99
              exported_namespace: giantswarm
              paused_value: true
            exp_annotations:
              description: "The clusters grumpy machinedeployment giantswarm/grumpy-def99 is paused."
              opsrecipe: capi-machinedeployment/
