rule_files:
  - capi-machinepool.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'capi_machinepool_spec_replicas{cluster_name="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinepool_status_replicas_available{cluster_name="clippaxy", name="clippaxy-def00", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinepool_spec_replicas{cluster_name="clippaxy", name="clippaxy-def99", exported_namespace="giantswarm"}'
        values: "0+3x75"
      - series: 'capi_machinepool_status_replicas_available{cluster_name="clippaxy", name="clippaxy-def99", exported_namespace="giantswarm"}'
        values: "0+2x75"
      - series: 'capi_machinepool_annotation_paused{paused_value="true",cluster_name="grumpy", name="grumpy-72r5c", exported_namespace="giantswarm"}'
        values: "0+1x75"
    alert_rule_test:
      - alertname: MachinePoolReplicasMismatch
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "false"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: clippaxy
              name: clippaxy-def99
              exported_namespace: giantswarm
            exp_annotations:
              description: "The clusters clippaxy machinepool giantswarm/clippaxy-def99 does not match the expected number of replicas for longer than 1h."
              opsrecipe: capi-machinepool/
      - alertname: MachinePoolPaused
        eval_time: 75m
        exp_alerts:
          - exp_labels:
              area: kaas
              cancel_if_outside_working_hours: "false"
              severity: notify
              team: phoenix
              topic: managementcluster
              cluster_name: grumpy
              name: grumpy-72r5c
              exported_namespace: giantswarm
              paused_value: true
            exp_annotations:
              description: "The clusters grumpy machinepool giantswarm/grumpy-72r5c is paused."
              opsrecipe: capi-machinepool/
