---
rule_files:
  - cert-manager.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'up{cluster_id="12345", cluster_type="workload_cluster", container="cert-manager", customer="giantswarm", installation="golem", instance="10.0.0.0:1234", job="12345-prometheus/workload-12345/0", namespace="kube-system", organization="giantswarm", pod="cert-manager-controller-7fcc585578-gnprd", provider="capa", service_priority="highest"}'
        values: "0+0x60"
    alert_rule_test:
      - alertname: CertManagerDown
        eval_time: 15m
        exp_alerts:
          - exp_labels:
              alertname: CertManagerDown
              area: platform
              cancel_if_kubelet_down: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: 12345
              cluster_type: workload_cluster
              container: cert-manager
              customer: giantswarm
              instance: 10.0.0.0:1234
              ip: 10.0.0.0
              job: 12345-prometheus/workload-12345/0
              namespace: kube-system
              organization: giantswarm
              pod: cert-manager-controller-7fcc585578-gnprd
              provider: capa
              installation: golem
              service_priority: highest
              severity: page
              team: shield
              topic: cert-manager
            exp_annotations:
              description: "cert-manager in namespace kube-system is down."
              runbook_url: https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/cert-manager-down/
  - interval: 1m
    input_series:
      - series: 'up{cluster_id="12345", cluster_type="workload_cluster", container="cert-manager", customer="giantswarm", installation="golem", instance="10.0.0.0:1234", job="12345-prometheus/workload-12345/0", namespace="kube-system", organization="giantswarm", pod="cert-manager-controller-7fcc585578-gnprd", provider="capa", service_priority="highest"}'
        values: "1+0x60"
    alert_rule_test:
      - alertname: CertManagerDown
        eval_time: 15m
