---
rule_files:
  - grafana.rules.yml

tests:
  - interval: 1m
    input_series:
      - series: 'up{service="grafana", cluster_id="gauss", cluster_type="management_cluster", installation="gauss", instance="grafana"}'
        values: "1+0x20 0+0x100"
    alert_rule_test:
      - alertname: GrafanaDown
        eval_time: 90m
        exp_alerts:
          - exp_labels:
              area: platform
              cancel_if_cluster_status_creating: "true"
              cancel_if_cluster_status_deleting: "true"
              cancel_if_cluster_status_updating: "true"
              cancel_if_outside_working_hours: "true"
              cluster_id: "gauss"
              cluster_type: management_cluster
              installation: "gauss"
              instance: "grafana"
              service: "grafana"
              severity: page
              team: atlas
              topic: observability
            exp_annotations:
              description: "Grafana (grafana) is down."
              runbook_url: "https://intranet.giantswarm.io/docs/support-and-ops/ops-recipes/grafana-down/"
              __dashboardUid__: qRQXmRnik
