---
rule_files:
  - inhibit.prometheus-agent.rules.yml

tests:
  - interval: 1m
    input_series:
      # - cluster 1: "clu01"
      - series: 'kube_namespace_created{cluster_id="myinstall", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", endpoint="http", installation="myinstall", instance="100.64.25.34:8080", job="kube-state-metrics", namespace="clu01-prometheus", node="ip-10-0-5-14.eu-central-1.compute.internal", organization="giantswarm", pipeline="testing", pod="prometheus-operator-app-kube-state-metrics-f7b868f49-ngvr8", service="prometheus-operator-app-kube-state-metrics"}'
        values: '1671707388+0x40'
      # - cluster 2: "clu02"
      - series: 'kube_namespace_created{cluster_id="myinstall", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", endpoint="http", installation="myinstall", instance="100.64.25.34:8080", job="kube-state-metrics", namespace="clu02-prometheus", node="ip-10-0-5-14.eu-central-1.compute.internal", organization="giantswarm", pipeline="stable", pod="prometheus-operator-app-kube-state-metrics-f7b868f49-ngvr8", service="prometheus-operator-app-kube-state-metrics"}'
        values: '1671707388+0x40'
      # - cluster 3: "myinstall", the install name
      - series: 'kube_namespace_created{cluster_id="myinstall", cluster_type="management_cluster", container="kube-state-metrics", customer="giantswarm", endpoint="http", installation="myinstall", instance="100.64.25.34:8080", job="kube-state-metrics", namespace="myinstall-prometheus", node="ip-10-0-5-14.eu-central-1.compute.internal", organization="giantswarm", pipeline="stable", pod="prometheus-operator-app-kube-state-metrics-f7b868f49-ngvr8", service="prometheus-operator-app-kube-state-metrics"}'
        values: "1671707388+0x40"
      # prometheus-agent app info for "clu01"
      - series: 'app_operator_app_info{name="prometheus-agent", app_version="2.40.5", catalog="giantswarm-playground", cluster_id="myinstall", cluster_missing="false", cluster_type="management_cluster", customer="giantswarm", deployed_version="0.1.7", endpoint="web", installation="myinstall", instance="app-exporter", job="app-exporter", namespace="clu01", node="ip-10-0-5-141.eu-central-1.compute.internal", organization="giantswarm", pipeline="stable", pod="app-exporter-6865c9c648-sg5vg", service="app-exporter", status="deployed", team="atlas", upgrade_available="false", version="0.1.7", version_mismatch="false"}'
        values: "1+0x40"
    alert_rule_test:
      #- alertname: InhibitionClusterIsNotRunningPrometheusAgent
      #  eval_time: 1m
      - alertname: InhibitionClusterIsNotRunningPrometheusAgent
        eval_time: 10m
        exp_alerts:
          - exp_labels:
              area: platform
              team: atlas
              topic: monitoring
              cluster_is_not_running_prometheus_agent: "true"
              cluster_id: "clu02"
            exp_annotations:
              description: "Cluster (clu02) is not running Prometheus Agent."
